# Caesar Cipher Assignment

```python
# preamble code, imports, etc...
```

## Section 

In this programming assignment, you will implement a simple Caesar encryption algorithm in Python. The Caesar cipher is one of the earliest and simplest encryption techniques, where each letter in the plaintext is shifted a certain number of places down or up the alphabet.

### Question

**Write a Python function `caesar_encrypt(plaintext)` that takes a plaintext string text as input and returns the encrypted ciphertext string using the Caesar cipher with a shift of 3 (to the right).**

_Ensure that your encryption function handles wrapping around the alphabet correctly. For example, shifting 'z' by 1 should result in 'a'. The input plaintext can be a mixture of uppercase and lower case letters. The output should be lowercase only._

_The input plaintext may contain spaces, periods (.) and commas(,). Your resulting ciphertext should ignore these three characters and output the plaintext as a sequence of letters without any spaces, periods, and commas._

#### Solution

<!-- grade_id: cell-15048376ca56a9a2, points: 0, grade: false, solution: true -->
```python
def caesar_encrypt(plaintext):
    shift = 3
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    # Create a dictionary for letter shifting
    shifted_alphabet = {alphabet[i]: alphabet[(i + shift) % 26] for i in range(len(alphabet))}
    
    # Prepare the ciphertext
    ciphertext = ''
    
    for char in plaintext:
        # Convert to lowercase
        char = char.lower()
        if char in alphabet:
            # Append the shifted character
            ciphertext += shifted_alphabet[char]
        # Ignore spaces, periods, and commas
        elif char in ' .,':
            continue
        else:
            # If character is not in alphabet, add it unchanged
            ciphertext += char
            
    return ciphertext
```

#### Auto-Grading
- Instructor Note: feel free to use any combination of these as you see fit and modify/remove the cells as needed

##### Visible Test Cases
- visible to learners in notebook
- learners can execute cell with tests while coding

<!-- grade_id: cell-f41cfd7e614545e1, points: 10, grade: true, solution: false -->
```python
### BEGIN TESTS
# Test basic functionality
assert caesar_encrypt("abc") == "def", "Error: Basic encryption not working"

# Test wrap-around (z to c)
assert caesar_encrypt("xyz") == "abc", "Error: Wrap-around not handled correctly"

# Test with uppercase characters (should be converted to lowercase)
assert caesar_encrypt("ABC") == "def", "Error: Uppercase characters not handled correctly"
### END TESTS
```

##### Hidden Tests
- will not be visible for the learner
- only run when assignment is submitted using submit button

<!-- grade_id: cell-3c2955b00b147106, points: 10, grade: true, solution: false -->
```python

```