# Playfair cipher programming assignment

```python
# preamble code, imports, etc...
```

## Section 

In this programming assignment, you will implement the Playfair cipher encryption algorithm in Python. 

### Question

**Write a Python function `playfair_encrypt(plaintext, key)` that takes a plaintext string text and a key as input and returns the encrypted ciphertext string using the Playfair cipher .**

_Ensure that your encryption function handles wrapping around the alphabet correctly. For example, shifting 'z' by 1 should result in 'a'. The input plaintext can be a mixture of uppercase and lower case letters. The output should be lowercase only._

_The input plaintext may contain spaces, periods (.) and commas(,). Your resulting ciphertext should ignore these three characters and output the plaintext as a sequence of letters without any spaces, periods, and commas._

#### Solution

<!-- grade_id: cell-a1b178dab5c74195, points: 0, grade: false, solution: true -->
```python
def playfair_encrypt(plaintext, key):

    def prepare_key(key):
        key = key.lower().replace('j', 'i')
        seen = set()
        key_square = []
        for char in key:
            if char not in seen and char.isalpha():
                seen.add(char)
                key_square.append(char)
        alphabet = 'abcdefghiklmnopqrstuvwxyz'
        for char in alphabet:
            if char not in seen:
                key_square.append(char)
        return key_square


    def create_digraphs(text):
        text = text.lower().replace(' ', '').replace('.', '').replace(',', '')
        text = text.replace('j', 'i')
        digraphs = []
        i = 0
        while i < len(text):
            if i + 1 < len(text) and text[i] == text[i + 1]:
                digraphs.append(text[i] + 'x')
                i += 1
            elif i + 1 < len(text):
                digraphs.append(text[i] + text[i + 1])
                i += 2
            else:
                digraphs.append(text[i] + 'x')
                i += 1
        return digraphs


    def find_position(char, key_square):
        index = key_square.index(char)
        row = index // 5
        col = index % 5
        return row, col


    def encrypt_pair(a, b, key_square):
        row_a, col_a = find_position(a, key_square)
        row_b, col_b = find_position(b, key_square)
        if row_a == row_b:
            return key_square[row_a * 5 + (col_a + 1) % 5] + key_square[row_b * 5 + (col_b + 1) % 5]
        elif col_a == col_b:
            return key_square[((row_a + 1) % 5) * 5 + col_a] + key_square[((row_b + 1) % 5) * 5 + col_b]
        else:
            return key_square[row_a * 5 + col_b] + key_square[row_b * 5 + col_a]


    key_square = prepare_key(key)
    digraphs = create_digraphs(plaintext)


    ciphertext = ""
    for digraph in digraphs:
        ciphertext += encrypt_pair(digraph[0], digraph[1], key_square)
    
    return ciphertext
```

#### Auto-Grading
- Instructor Note: feel free to use any combination of these as you see fit and modify/remove the cells as needed

##### Visible Test Cases
- visible to learners in notebook
- learners can execute cell with tests while coding

<!-- grade_id: cell-85d32c0776b9cd10, points: 25, grade: true, solution: false -->
```python
### BEGIN TESTS
# Test basic functionality
assert playfair_encrypt("juliuscaesarisdead", "earth") == "nomkoukctortmobrrc", "Error: Basic encryption not working"

### END TESTS
```

##### Hidden Tests
- will not be visible for the learner
- only run when assignment is submitted using submit button

<!-- grade_id: cell-3d6211e5c8a90790, points: 25, grade: true, solution: false -->
```python

```

```python

```